#### Bi 410/510 (Fall 2019)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

When you are finished submit this notebook through Canvas.

# <span style="color:steelblue;">Milestone 4a</span>

### <span style="color:steelblue;">GC Content</span>

Implement a function named `gc_content` that computes the GC content of a DNA sequence (the percentage of bases that are either Gs or Cs), rounded to two significant digits. 

**Do not use the `count` method.** Instead, write a single for loop that checks each letter in the input string and updates a counter if a letter is a G or a C. 

Example:
```
>>> gc('GATTACA')
28.57
```

<!-- grade_id: gc_4a, points: 3, grade: true, solution: true -->
```python
def gc(seq):
    count=0
    for i in range(len(seq)):
        if (seq[i]=="G" or seq[i]=="C"):
            count+=1 
    percentage=round((count/len(seq))*100,2)
    return percentage
```

##### Tests 

Use this code cell to test your function as you write it.

```python
gc("GATTACA")
```

Output:
```
28.57
```

We'll use these code cells to check your code.

<!-- grade_id: cell-425fdb370cea0bad, points: 1, grade: true, solution: false -->
```python
assert gc('GATTACA') == 28.57
```

<!-- grade_id: cell-4bd98d4a857264c3, points: 1, grade: true, solution: false -->
```python
assert gc('AAACCCGGGTTT') == 50.0
```

<!-- grade_id: cell-7529958e880100b4, points: 1, grade: true, solution: false -->
```python
assert gc('AAATTT') == 0.0
```

### <span style="color:steelblue;">Acronym</span>

Implement a function named `acronym` that will create an acronym from a list of words.  The result should be a single string that has the first character of each word from the list.

Example:
```
>>> acronym(["North", "Atlantic", "Treaty", "Organization"])
'NATO'
```

<!-- grade_id: acronym, points: 3, grade: true, solution: true -->
```python
def acronym(lst):
    # YOUR CODE HERE
    acronym=''
    for i in range(len(lst)):
        word=lst[i]
        acronym+=word[0]
    return acronym
```

##### Tests 

Use this code cell to test your function as you write it.

```python
acronym(["North", "Atlantic", "Treaty", "Organization"])
```

Output:
```
'NATO'
```

We'll use these code cells to check your code.

<!-- grade_id: cell-996686965e92ca7e, points: 1, grade: true, solution: false -->
```python
assert acronym(["North", "Atlantic", "Treaty", "Organization"]) == 'NATO'
```

<!-- grade_id: cell-cdb3e7cce1cffac3, points: 1, grade: true, solution: false -->
```python
assert acronym(["Martin", "Luther", "King"]) == 'MLK'
```

### <span style="color:steelblue;">Birthday Paradox</span>

A fun ice-breaker for a new group of people is to have everyone compare birthdays to see if any two people were born on the same day of the year. It may not seem intuitive, but with only 23 people in a group there is a 50\% chance that two of them will have the same birthday. With a group of 100 the probability that two people will have the same birthday is more than 99.999\%.

To simulate this phenomenon in Python write a function named `paradox`. Create an empty list, and then use a loop that adds numbers to the list. On each iteration:
* choose a random number between 1 and 365 (use the `randint` function from the `random` library)
* if the number is already in the list exit the loop
* otherwise add the number to the list and continue.

The function should return the number of times the loop iterates before finding a duplicate. If it will help, you can print the contents of your list and the random number you draw on each loop iteration. Here is a sample output with this more verbose output:
```
>>> paradox()
look for 343 in []
look for 130 in [343]
look for 271 in [343, 130]
look for 192 in [343, 130, 271]
look for 97 in [343, 130, 271, 192]
look for 322 in [343, 130, 271, 192, 97]
look for 271 in [343, 130, 271, 192, 97, 322]
6
```

<!-- grade_id: paradox, points: 3, grade: true, solution: true -->
```python
from random import randint

def paradox():
    # YOUR CODE HERE
    birthdays=[]
    num=randint(1,365)
    print("look for ", num,"in ",birthdays)
    birthdays=[num]
    for i in range(365):
        num=randint(1,365)
        print("look for ", num,"in ",birthdays)
        if not num in birthdays:
            birthdays.append(num)
        else:
            return(i+1)
```

##### Tests 

Use this code cell to test your function as you write it.

```python
paradox()
```

Output:
```
look for  347 in  []
look for  251 in  [347]
look for  18 in  [347, 251]
look for  331 in  [347, 251, 18]
look for  41 in  [347, 251, 18, 331]
look for  31 in  [347, 251, 18, 331, 41]
look for  315 in  [347, 251, 18, 331, 41, 31]
look for  110 in  [347, 251, 18, 331, 41, 31, 315]
look for  221 in  [347, 251, 18, 331, 41, 31, 315, 110]
look for  199 in  [347, 251, 18, 331, 41, 31, 315, 110, 221]
look for  96 in  [347, 251, 18, 331, 41, 31, 315, 110, 221, 199]
look for  129 in  [347, 251, 18, 331, 41, 31, 315, 110, 221, 199, 96]
look for  41 in  [347, 251, 18, 331, 41, 31, 315, 110, 221, 199, 96, 129]
12
```

We'll use these code cells to check your code.  Each time you call your function you should get back a number between 1 and 365.

<!-- grade_id: cell-3a430d9c6a27297d, points: 2, grade: true, solution: false -->
```python
for i in range(5):
    assert 1 <= paradox() <= 365
```

Output:
```
look for  339 in  []
look for  122 in  [339]
look for  8 in  [339, 122]
look for  75 in  [339, 122, 8]
look for  314 in  [339, 122, 8, 75]
look for  22 in  [339, 122, 8, 75, 314]
look for  352 in  [339, 122, 8, 75, 314, 22]
look for  162 in  [339, 122, 8, 75, 314, 22, 352]
look for  252 in  [339, 122, 8, 75, 314, 22, 352, 162]
look for  134 in  [339, 122, 8, 75, 314, 22, 352, 162, 252]
look for  347 in  [339, 122, 8, 75, 314, 22, 352, 162, 252, 134]
look for  238 in  [339, 122, 8, 75, 314, 22, 352, 162, 252, 134, 347]
look for  144 in  [339, 122, 8, 75, 314, 22, 352, 162, 252, 134, 347, 238]
look for  171 in  [339, 122, 8, 75, 314, 22, 352, 162, 252, 134, 347, 238, 144]
look for  358 in  [339, 122, 8, 75, 314, 22, 352, 162, 252, 134, 347, 238, 144, 171]
look for  129 in  [339, 122, 8, 75, 314, 22, 352, 162, 252, 134, 347, 238, 144, 171, 358]
look for  262 in  [339, 122, 8, 75, 314, 22, 352, 162, 252, 134, 347, 238, 144, 171, 358, 129]
look for  347 in  [339, 122, 8, 75, 314, 22, 352, 1
```

### <span style="color:red;">Short Answer Question</span>

Don't forget to complete the short-answer question on the exam page.

Double-click on the markdown cell below and enter your answer in this cell.

<!-- grade_id: cell-cd72593eb7ea9d6f, points: 4, grade: true, solution: true -->
1- The for loop is executed 5 times<br>
2- char is assigned 1<br>
3- msg is assigned I30<br>